# Algorithme des k plus proches voisins
l’algorithme des k plus proches voisins (K-NN où K-nearest neighbors) est une méthode utilisée pour la classification de données. Son fonctionnement peut être assimilé à l’analogie suivante *dis moi qui sont tes voisins, je te dirais qui tu es…*.

Dans cette activité, vous allez travailler sur un jeu de données fictives : 30 points sont disposés sur un plateau de taille $100 \times 100$. Ces points sont de couleur *rouge*, *verte*, *bleue* ou *noire*. 

Le problème est le suivant : étant données les coordonnées d'un point choisi au hasard sur notre plateau, peut-on deviner sa couleur ? La réponse à cette question sera apportée par l'algorithme des k plus proches voisins que vous devrez réaliser en suivant les indications ci-dessous.

```python
# Import des bibliothèques nécessaires

import matplotlib.pyplot as plt
import numpy as np
from PIL import Image
from random import randint
from math import sqrt
import hashlib
import sys
%matplotlib inline
```

```python
# Constantes
MAX_RANGE = 100                        # les coordonnées sont entre 0 et 99
PRECISION = 10*sys.float_info.epsilon  # precision des tests sur les flottants
COULEURS = {'red':(255,0,0), 'green':(0,255,0), 'blue':(0,0,255), 'black':(0,0,0)}
FORMES = {"red":"o", "blue":"^", "green":"s", "black":"x"}
```

## 1. Jeu de données

les informations de départ sont données sous forme d'une liste de tuple de trois termes
   - le premier est l'abscisse
   - le deuxième est l'ordonnée 
   - le troisième la couleur: red blue, green et black

### Récupération d'un jeu existant
Validez les 3 cellules ci-dessous et étudiez la manière dont les données sont représentées dans le dictionnaire ***data***

```python
data=[(52, 2, 'black'), (59, 0, 'black'), (68, 8, 'black'), (77, 11, 'black'), (28, 21, 'black'), (0, 66, 'blue'),
      (0, 77, 'blue'), (28, 61, 'blue'), (21, 56, 'blue'), (29, 66, 'blue'), (10, 40, 'blue'), (10, 81, 'blue'), 
      (30, 49, 'blue'), (83, 87, 'green'), (96, 56, 'green'), (45, 77, 'green'), (73, 69, 'green'), 
      (53, 77, 'green'), (34, 68, 'green'), (70, 63, 'green'), (85, 34, 'red'), (36, 54, 'red'), (33, 54, 'red'),
      (65, 52, 'red'), (45, 45, 'red'), (70, 62, 'red'), (48, 66, 'red'), (63, 28, 'red'), (40, 28, 'red'),
      (82, 18, 'red')]
```

```python
def pts_couleur(couleur):
    liste=[]
    for el in data:
        if el[2]==couleur:
            liste.append((el[0],el[1]))
    return liste

ptsNoirs=pts_couleur('black')
```

```python
print("ce jeu contient ", len(data), " données")
print("les points noirs sont ", ptsNoirs)
```

### Représentation graphique du jeu de données

Pour visualiser notre jeu de données, chaque donnée sera un point dont les coordonnées seront les deux valeurs entre 0 et 100 et dont la couleur et la forme corrspondront au critère associé. Pour simplifier la lecture sur un document noir et blanc, nous adopterons la convention suivante :
- le rouge sera représenté par un point rond
- le vert sera représenté par un carré
- le bleu sera représenté par un triangle
- le noir sera représenté par une croix

```python
def dessine_points(data):
    plt.rcParams["figure.figsize"] = (4, 4)
    plt.axis('equal')
    for el in data:
        plt.scatter(el[0], el[1], c=el[2], marker=FORMES[el[2]])
    plt.show()
dessine_points(data)
```

## 2.Implémentation de l'algorithme des k plus proches voisins

On cherche dans cette partie à écrire un programme qui, étant données les coordonnées d'un point, devine la couleur associée en utilisant l'algorithme des k plus proches voisins.

### a) Calcul de la distance entre deux points
Ecrire une fonction **distance** prenant en paramètre deux points (un point sera représenté par un tuple (x,y)) et renvoyant la distance entre ces deux points.

<!-- grade_id: cell-832cd6f12ac11590, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def distance(pointA, pointB):
    xA, yA = pointA
    xB, yB = pointB
    return sqrt((xB-xA)**2 + (yB-yA)**2)
### END SOLUTION
```

<!-- grade_id: cell-8ad51ee90553e693, points: 1, grade: true, solution: false -->
```python
# test de la fonction
cible = (5, 5)
voisin1 = (3, 4)
voisin2 = (0, 0)
print(distance(cible, voisin1))
print(distance(cible, voisin2))

# Verification automatisee
assert abs(distance(cible, voisin1) - 2.23606797749979) < PRECISION
assert abs(distance(cible, voisin2) - 7.0710678118654755) < PRECISION
### BEGIN HIDDEN TESTS
assert abs(distance((34,52), (12, 30)) - 31.11269837220809) < PRECISION
assert abs(distance((34,52), (34, 52)) - 0.0) < PRECISION
### END HIDDEN TESTS
```

### b) Calcul des distances de tous les points à une origine donnée

Construire une fonction ***k_proches_voisins*** qui prend en arguments
- le jeu de données
- les coordonnées de la cible: le point dont on cherche à déterminer a nature
- la valeur de *k* : nombre de plus proches voisins à considérer (il faut 1 par défaut)

et renvoyant la liste des k plus proches voisins sous forme d'un tuple (distance, propriete)

<!-- grade_id: cell-1de10ed9d3c42934, points: 0, grade: false, solution: true -->
```python
def k_plus_proches(data,cible,k=1):
    liste_voisin=[]
    for el in data:
        d=distance(cible,(el[0],el[1]))
        liste_voisin.append((d,el[2]))
    liste_voisin.sort()
    return liste_voisin[:k]
```

<!-- grade_id: cell-877252982aaabda8, points: 3, grade: true, solution: false -->
```python
# test de la fonction
reponse = k_plus_proches(data, (5,5), 5)
for r in reponse:
    print(r)
# tests automatises
assert abs(reponse[0][0]-28.0178514522438) < PRECISION
assert reponse[0][1][2] == 'black'
assert abs(reponse[1][0]-35.35533905932738) < PRECISION
assert reponse[1][1][2] == 'blue'
assert abs(reponse[2][0]-41.88078318274385) < PRECISION
assert reponse[2][1][2] == 'red'
assert abs(reponse[3][0]-47.095647357266465) < PRECISION
assert reponse[3][1][2] == 'black'
assert abs(reponse[4][0]-50.60632371551998) < PRECISION
assert reponse[4][1][2] == 'blue'
### BEGIN HIDDEN TESTS
reponse1 = k_plus_proches(data, (51,25), 2)
assert abs(reponse1[0][0]-11.40175425099138) < PRECISION
assert reponse1[0][1][2] == 'red'
assert abs(reponse1[1][0]-12.36931687685298) < PRECISION
assert reponse1[1][1][2] == 'red'
### END HIDDEN TESTS
```

### c) Prédiction de la valeur du caractère en fonction des k plus proches voisins

Déduire de la fonction précédante la fonction ***categorie_devine*** qui prend en arguments
- le jeu de données
- les coordonnées du point origine dont on cherche à déterminer la nature
- la valeur de $k$ : nombre de plus proches voisins à considérer

qui complète  un dictionnaire dont les clés sont les couleurs et les valeurs une liste de deux éléments: 
- le nombre de fois où la couleur apparait dans la liste des k plus proches voisins
- la distance cumulée des points d'une même couleur à cible

et qui renvoie la couleur de la cible, à savoir la couleur la plus fréquente parmi les $k$ plus proches voisins. 

En cas d'égalité, on choisira la couleur dont la **somme des distances correspondante est la plus petite**.


<!-- grade_id: cell-f4566c3bd94a823c, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def categorie_devine(data, origine, k=1):
    dico = {}
    kpv = k_plus_proches(data, origine, k)
    #complétez le code 
        
    return coulmax
### END SOLUTION
```

<!-- grade_id: cell-e55f693ff4c3955f, points: 2, grade: true, solution: false -->
```python
# test de la fonction
reponse = categorie_devine(data, (5,5),5)
print(reponse)
# test automatise
assert categorie_devine(data, ( 51 , 42 ),6) == 'red'
assert categorie_devine(data, ( 40 , 40 ),3) == 'red'
assert categorie_devine(data, ( 40 , 60 ),3) == 'red'
assert categorie_devine(data, ( 40 , 80 ),3) == 'green'
assert categorie_devine(data, ( 60 , 20 ),3) == 'black'
assert categorie_devine(data, ( 60 , 40 ),30) == 'red'
### BEGIN HIDDEN TESTS
assert categorie_devine(data, ( 20 , 20 ),3) == 'black'
assert categorie_devine(data, ( 20 , 40 ),3) == 'blue'
assert categorie_devine(data, ( 20 , 60 ),3) == 'blue'
assert categorie_devine(data, ( 20 , 80 ),3) == 'blue'
assert categorie_devine(data, ( 40 , 20 ),3) == 'black'
assert categorie_devine(data, ( 60 , 60 ),3) == 'red'
assert categorie_devine(data, ( 60 , 80 ),3) == 'green'
assert categorie_devine(data, ( 80 , 20 ),3) == 'red'
assert categorie_devine(data, ( 80 , 40 ),3) == 'red'
assert categorie_devine(data, ( 80 , 60 ),3) == 'green'
assert categorie_devine(data, ( 80 , 80 ),3) == 'green'
### END HIDDEN TESTS
```

## 3. Régionnement du plan

En partant de la représentation graphique de notre jeu de données, nous allons colorier chaque point du plan avec la couleur prédite par l'algorithme des k plus proches voisins. Nous visualiserons l'influence du paramètre $k$ dans la prédiction obtenue.

L'idée dans cette partie est la suivante : on parcourt tous les points du plan et on colorie le pixel avec la couleur prédite par l'algorithme en fonction de ses k plus proches voisins.

Pour cette tâche, nous aurons besoin de savoir comment créer une image pixels par pixels. Etudiez l'exemple ci-dessous qui vous montrera l'utilisation de la librairie python *PIL* dédiée à cela.

```python
# On cree une image vide
monimage=Image.new("RGB",(MAX_RANGE, MAX_RANGE))

# On definit une couleur selon les valeurs Rouge, Vert et Bleu entre 0 et 255
R,V,B = (0,64,21)

# On calcule une image de 100 pixels par 100 pixels
for x in range(MAX_RANGE):
    for y in range(MAX_RANGE):
        monimage.putpixel((x,MAX_RANGE-y-1),(R,V,B))

# On zoome sur l'image pour lui donner la taille demandée
monimage.resize((240,240))
```

L'image ci-dessus vous semble trop monotonne ? amusez-vous à modifier la couleur des pixels en fonction des coordonnées. Par exemple : mettez $2x$ pour le canal vert et $2y$ pour le canal bleu.

**Remarque** : On rappelle que l'origine du repère dans *PIL* est le point en haut à gauche et l'axe des ordonnées est orienté vers le bas.
Dans *matplotlib*, l'origine du repère est le point en bas à gauche et l'axe des ordonnées est orienté vers le haut. 
On tiendra compte de ces spécificités de manière à ce que le régionnement obtenu corresponde à la représentation des points obtenue ci-dessus.


### Régionnement du plan

Ecrire une fonction ***remplissage*** qui prend en paramètre le jeu de données et la valeur de $k$ et qui renvoie une image du plan colorié avec une résolution de 240x240 pixels.

**Indication** : On pourra utiliser le dictionnaire **COULEURS** donné en constante au début de ce document pur convertir les propriétés de couleur en valeurs au format (R,V,B).

**Attention** : Penser à renvoyer l'image fabriquée (avec *return*) avec les dimensions demandées à la fin de la fonction ***remplissage*** !

<!-- grade_id: cell-fdbe2eef705fadde, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def remplissage(data, k):
    monimage=Image.new("RGB",(MAX_RANGE, MAX_RANGE))
    #complétez le code
    
    return monimage.resize((240,240))
### END SOLUTION
```

<!-- grade_id: cell-bcfb90631ada15e6, points: 2, grade: true, solution: false -->
```python
# test automatise
monImage = remplissage(data, 3)
assert monImage.size == (240,240)
assert type(monImage) == Image.Image

monImage
```

Exécutez cette fonction avec plusieurs valeurs de $k$ variant entre 4 et 10

```python
remplissage(data, 4)
```

```python
remplissage(data, 5)
```

```python
remplissage(data, 6)
```

```python
remplissage(data, 7)
```

```python
remplissage(data, 8)
```

```python
remplissage(data, 9)
```

```python
remplissage(data, 10)
```